Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All).

Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE", as well as your name and collaborators below:

```python
NAME = ""
COLLABORATORS = ""
```

---

# Continuous Optimization (Chapter 7: MML Book)

Latex macros
$$
\newcommand{\calA}{{\cal A}}
\newcommand{\calB}{{\cal B}}
\newcommand{\calC}{{\cal C}}
\newcommand{\calD}{{\cal D}}
\newcommand{\calE}{{\cal E}}
\newcommand{\calF}{{\cal F}}
\newcommand{\calG}{{\cal G}}
\newcommand{\calH}{{\cal H}}
\newcommand{\calI}{{\cal I}}
\newcommand{\calJ}{{\cal J}}
\newcommand{\calK}{{\cal K}}
\newcommand{\calL}{{\cal L}}
\newcommand{\calM}{{\cal M}}
\newcommand{\calN}{{\cal N}}
\newcommand{\calO}{{\cal O}}
\newcommand{\calP}{{\cal P}}
\newcommand{\calQ}{{\cal Q}}
\newcommand{\calR}{{\cal R}}
\newcommand{\calS}{{\cal S}}
\newcommand{\calT}{{\cal T}}
\newcommand{\calU}{{\cal U}}
\newcommand{\calV}{{\cal V}}
\newcommand{\calW}{{\cal W}}
\newcommand{\calX}{{\cal X}}
\newcommand{\calY}{{\cal Y}}
\newcommand{\calZ}{{\cal Z}}
\newcommand{\setA}{\textsf{A}}
\newcommand{\setB}{\textsf{B}}
\newcommand{\setC}{\textsf{C}}
\newcommand{\setD}{\textsf{D}}
\newcommand{\setE}{\textsf{E}}
\newcommand{\setF}{\textsf{F}}
\newcommand{\setG}{\textsf{G}}
\newcommand{\setH}{\textsf{H}}
\newcommand{\setI}{\textsf{I}}
\newcommand{\setJ}{\textsf{J}}
\newcommand{\setK}{\textsf{K}}
\newcommand{\setL}{\textsf{L}}
\newcommand{\setM}{\textsf{M}}
\newcommand{\setN}{\textsf{N}}
\newcommand{\setO}{\textsf{O}}
\newcommand{\setP}{\textsf{P}}
\newcommand{\setQ}{\textsf{Q}}
\newcommand{\setR}{\textsf{R}}
\newcommand{\setS}{\textsf{S}}
\newcommand{\setT}{\textsf{T}}
\newcommand{\setU}{\textsf{U}}
\newcommand{\setV}{\textsf{V}}
\newcommand{\setW}{\textsf{W}}
\newcommand{\setX}{\textsf{X}}
\newcommand{\setY}{\textsf{Y}}
\newcommand{\setZ}{\textsf{Z}}
\newcommand{\bfa}{\mathbf{a}}
\newcommand{\bfb}{\mathbf{b}}
\newcommand{\bfc}{\mathbf{c}}
\newcommand{\bfd}{\mathbf{d}}
\newcommand{\bfe}{\mathbf{e}}
\newcommand{\bff}{\mathbf{f}}
\newcommand{\bfg}{\mathbf{g}}
\newcommand{\bfh}{\mathbf{h}}
\newcommand{\bfi}{\mathbf{i}}
\newcommand{\bfj}{\mathbf{j}}
\newcommand{\bfk}{\mathbf{k}}
\newcommand{\bfl}{\mathbf{l}}
\newcommand{\bfm}{\mathbf{m}}
\newcommand{\bfn}{\mathbf{n}}
\newcommand{\bfo}{\mathbf{o}}
\newcommand{\bfp}{\mathbf{p}}
\newcommand{\bfq}{\mathbf{q}}
\newcommand{\bfr}{\mathbf{r}}
\newcommand{\bfs}{\mathbf{s}}
\newcommand{\bft}{\mathbf{t}}
\newcommand{\bfu}{\mathbf{u}}
\newcommand{\bfv}{\mathbf{v}}
\newcommand{\bfw}{\mathbf{w}}
\newcommand{\bfx}{\mathbf{x}}
\newcommand{\bfy}{\mathbf{y}}
\newcommand{\bfz}{\mathbf{z}}
\newcommand{\bfalpha}{\boldsymbol{\alpha}}
\newcommand{\bfbeta}{\boldsymbol{\beta}}
\newcommand{\bfgamma}{\boldsymbol{\gamma}}
\newcommand{\bfdelta}{\boldsymbol{\delta}}
\newcommand{\bfepsilon}{\boldsymbol{\epsilon}}
\newcommand{\bfzeta}{\boldsymbol{\zeta}}
\newcommand{\bfeta}{\boldsymbol{\eta}}
\newcommand{\bftheta}{\boldsymbol{\theta}}
\newcommand{\bfiota}{\boldsymbol{\iota}}
\newcommand{\bfkappa}{\boldsymbol{\kappa}}
\newcommand{\bflambda}{\boldsymbol{\lambda}}
\newcommand{\bfmu}{\boldsymbol{\mu}}
\newcommand{\bfnu}{\boldsymbol{\nu}}
\newcommand{\bfomicron}{\boldsymbol{\omicron}}
\newcommand{\bfpi}{\boldsymbol{\pi}}
\newcommand{\bfrho}{\boldsymbol{\rho}}
\newcommand{\bfsigma}{\boldsymbol{\sigma}}
\newcommand{\bftau}{\boldsymbol{\tau}}
\newcommand{\bfupsilon}{\boldsymbol{\upsilon}}
\newcommand{\bfphi}{\boldsymbol{\phi}}
\newcommand{\bfchi}{\boldsymbol{\chi}}
\newcommand{\bfpsi}{\boldsymbol{\psi}}
\newcommand{\bfomega}{\boldsymbol{\omega}}
\newcommand{\bfxi}{\boldsymbol{\xi}}
\newcommand{\bfell}{\boldsymbol{\ell}}
\newcommand{\bfA}{\mathbf{A}}
\newcommand{\bfB}{\mathbf{B}}
\newcommand{\bfC}{\mathbf{C}}
\newcommand{\bfD}{\mathbf{D}}
\newcommand{\bfE}{\mathbf{E}}
\newcommand{\bfF}{\mathbf{F}}
\newcommand{\bfG}{\mathbf{G}}
\newcommand{\bfH}{\mathbf{H}}
\newcommand{\bfI}{\mathbf{I}}
\newcommand{\bfJ}{\mathbf{J}}
\newcommand{\bfK}{\mathbf{K}}
\newcommand{\bfL}{\mathbf{L}}
\newcommand{\bfM}{\mathbf{M}}
\newcommand{\bfN}{\mathbf{N}}
\newcommand{\bfO}{\mathbf{O}}
\newcommand{\bfP}{\mathbf{P}}
\newcommand{\bfQ}{\mathbf{Q}}
\newcommand{\bfR}{\mathbf{R}}
\newcommand{\bfS}{\mathbf{S}}
\newcommand{\bfT}{\mathbf{T}}
\newcommand{\bfU}{\mathbf{U}}
\newcommand{\bfV}{\mathbf{V}}
\newcommand{\bfW}{\mathbf{W}}
\newcommand{\bfX}{\mathbf{X}}
\newcommand{\bfY}{\mathbf{Y}}
\newcommand{\bfZ}{\mathbf{Z}}
\newcommand{\bfGamma}{\boldsymbol{\Gamma}}
\newcommand{\bfDelta}{\boldsymbol{\Delta}}
\newcommand{\bfTheta}{\boldsymbol{\Theta}}
\newcommand{\bfLambda}{\boldsymbol{\Lambda}}
\newcommand{\bfPi}{\boldsymbol{\Pi}}
\newcommand{\bfSigma}{\boldsymbol{\Sigma}}
\newcommand{\bfUpsilon}{\boldsymbol{\Upsilon}}
\newcommand{\bfPhi}{\boldsymbol{\Phi}}
\newcommand{\bfPsi}{\boldsymbol{\Psi}}
\newcommand{\bfOmega}{\boldsymbol{\Omega}}
\newcommand{\bbA}{\mathbb{A}}
\newcommand{\bbB}{\mathbb{B}}
\newcommand{\bbC}{\mathbb{C}}
\newcommand{\bbD}{\mathbb{D}}
\newcommand{\bbE}{\mathbb{E}}
\newcommand{\bbF}{\mathbb{F}}
\newcommand{\bbG}{\mathbb{G}}
\newcommand{\bbH}{\mathbb{H}}
\newcommand{\bbI}{\mathbb{I}}
\newcommand{\bbJ}{\mathbb{J}}
\newcommand{\bbK}{\mathbb{K}}
\newcommand{\bbL}{\mathbb{L}}
\newcommand{\bbM}{\mathbb{M}}
\newcommand{\bbN}{\mathbb{N}}
\newcommand{\bbO}{\mathbb{O}}
\newcommand{\bbP}{\mathbb{P}}
\newcommand{\bbQ}{\mathbb{Q}}
\newcommand{\bbR}{\mathbb{R}}
\newcommand{\bbS}{\mathbb{S}}
\newcommand{\bbT}{\mathbb{T}}
\newcommand{\bbU}{\mathbb{U}}
\newcommand{\bbV}{\mathbb{V}}
\newcommand{\bbW}{\mathbb{W}}
\newcommand{\bbX}{\mathbb{X}}
\newcommand{\bbY}{\mathbb{Y}}
\newcommand{\bbZ}{\mathbb{Z}}
$$



### Recall geometry of a derivative

![](imgs/geometry-of-derivative.png)

![](imgs/minimas-of-non-convex-function.png)

## Minimizing general functions

We cannot minimize general functions by solving

$$ \frac{\p f(\bfx)}{\p \bfx} = \mathbf{0}^\top$$ 

because the equation might not have a formula for it.

Instead we use iterative methods like gradient descent minimize general function $f(\bfx)$.

#### Definition (Directional derivative)

$\newcommand{\p}{\partial}$

Directional dervivative of a function $f(\bfx) : \bbR^n \to \bbR$ with respect to a given vector $\bfu \in \bbR$ is defined as

$$ D_\bfu f(\bfx) = \lim_{\epsilon \to 0} \frac{f(\bfx + \epsilon \bfu) - f(\bfx)}{\epsilon} $$

[Ref Khan Academy](https://www.khanacademy.org/math/multivariable-calculus/multivariable-derivatives/partial-derivative-and-gradient-articles/a/directional-derivative-introduction)

[Ref Libretexts](https://math.libretexts.org/Bookshelves/Calculus/Calculus_3e_(Apex)/12%3A_Functions_of_Several_Variables/12.06%3A_Directional_Derivatives)


##### Vector calculus chain rule (a theorem)

Given a function composition $\bff(\bfx) = \bfg(\bfh(\bfx)) = (\bfg \circ \bfh)(\bfx)$ where $\bfh : \bbR^n \to \bbR^m$, $\bfg : \bbR^m \to \bbR^p$ and $\bfh: \bbR^m \to \bbR^p$

$$ \frac{\p \bff}{\p \bfx} = \frac{\p \bff}{\p \bfh}\frac{\p \bfh}{\p \bfx} $$
or denoting the derivatives as Jacobian matrices we have,
$$ \calJ_\bfx \bff = \calJ_\bfh[\bff] \calJ_\bfx [ \bfh ] $$

##### Theorem (Directional derivative is gradient dot product with the direction)

Express the trajectory in the direction $\bfu$ as a function of time $t$ as
$$\bfg(t) = \bfx + t \bfu $$
Note that the Jacobian of $\bfg(t)$ wrt $t$ is simply $\bfu$,
$$\calJ_t \bfg(t) = \bfu$$
Recall the definition of directional derivative,
$$ D_\bfu f(\bfx) = \lim_{\epsilon \to 0} \frac{f(\bfx + \epsilon \bfu) - f(\bfx)}{\epsilon}. $$

Compare it with the derivative of $f(\bfg(t))$ with respect to $t$ at $t=0$
$$ \frac{\p f(\bfg(t))}{\p t} = \lim_{\epsilon \to 0} \frac{f(\bfx + (t + \epsilon) \bfu) - f(\bfx + t\bfu)}{\epsilon} \Big\vert_{t = 0}.$$
$$ \frac{\p f(\bfg(t))}{\p t} = \lim_{\epsilon \to 0} \frac{f(\bfx + \epsilon \bfu) - f(\bfx )}{\epsilon} = D_\bfu f(\bfx).$$

We can compute $\frac{\p f(\bfg(t))}{\p t}$ by chain rule,

$$  D_\bfu f(\bfx) = \calJ_t f(\bfg(t)) = \calJ_\bfx f(\bfx) \calJ_t \bfg = \nabla_\bfx^\top f(\bfx) \bfu$$

#### Theorem : The direction of stepest ascent and descent

Let $\hat{\bfu}$ be of unit magnitude. The directional derivative represents how the function changes in the direction $\hat{\bfu}$.

$$ D_{\hat{\bfu}} f(\bfx) =  \nabla_\bfx^\top f(\bfx) \hat{\bfu} = \|\nabla_\bfx f(\bfx)\|\cos(\theta),$$ 
where $\theta$ is the angle between $\nabla_\bfx f(\bfx)$ and $\hat{\bfu}$. The change is maximum when $\theta = 0$ and $\cos(\theta) = 1$ and the change is minimum when $\theta = 180^\circ$ and $\cos(\theta) = -1$. 

In other words the function $f$ increases the most (stepest ascent) when $\hat{\bfu}  \propto \nabla_\bfx f(\bfx)$ and decreases the most (steepest descent) when $\hat{\bfu} \propto - \nabla_\bfx f(\bfx)$.

### Gradient descent method
(Section 9.3 of [Convex Optimization by Stephen Boyd and Lieven Vandenberghe](https://web.stanford.edu/~boyd/cvxbook/bv_cvxbook.pdf))


1. Start from a random point $\bfx_0$, $\bfx_t \leftarrow \bfx_0$.
2. Move in the direction opposite to $\nabla_\bfx f(\bfx)$. If we were at $\bfx_t$, then the next point is at $\bfx_{t+1} = \bfx_t - \alpha_t \nabla_\bfx f(\bfx)$, where $\alpha_t > 0$ is a positive scalar, called the step size or the learning rate.
3. Stop when the gradient is almost zero $\|\nabla_\bfx f(\bfx)\| < 10^{-4}$.

This corresponds to the following algorithm:

$\bfx_t = \bfx_0$<br/>
while ($\|\nabla_\bfx f(\bfx)\| > 10^{-4}$) {
    $\bfx_t \leftarrow \bfx_t - \alpha_t \nabla_\bfx f(\bfx)$
}


Algorithm 9.3 Gradient descent method.

**given** a starting point $\bfx \in \text{dom}{f}$ .

**repeat**
1. $\Delta x = -\nabla f(\bfx)$.
2. Choose step size $\alpha$ 
3. Update. $\bfx := \bfx + \alpha \bfx$

**until** stopping criterion is satisfied.

![](https://upload.wikimedia.org/wikipedia/commons/f/ff/Gradient_descent.svg)


<video src="https://upload.wikimedia.org/wikipedia/commons/transcoded/4/4c/Gradient_Descent_in_2D.webm/Gradient_Descent_in_2D.webm.720p.vp9.webm"  controls>
	This is fallback content to display for user agents that do not support the video tag.
</video>

## Gradient visualization

```python
import matplotlib.pyplot as plt
import numpy as np
from matplotlib import animation, rc
rc('animation', html='jshtml')
def plot_gradients(func, gradfunc):
    x, y = np.mgrid[-3:3:21j,
                    -3:3:21j]
    bfx = np.array([x, y])
    f = func(x,y)
    [dfdx, dfdy] = gradfunc(x,y)
    fig, ax = plt.subplots()
    ctr = ax.contour(x, y, f, 20, cmap='Blues_r')
    ax.quiver(x, y, dfdx, dfdy)
    ax.plot([0], [0], 'ro') 
    ax.text(0, 0, '$x^*$', color='r')
    ax.set_xlabel('$x$')
    ax.set_ylabel('$y$')
    plt.show()
    
```

```python
def f(x, y): return  x**2 + y**2
def gradf(x, y): return 2*x, 2*y
plot_gradients(f, gradf)
```

### Quadratic function example:

A quadratic function $f(\bfx) = x_1^2 + x_2^2$ has level sets as circles:
$$ S(f, c) = \{\bfx : x_1^2 + x_2^2 = c\} $$

It has the parameteric form as
$$ S(f, c) = \{\bfg(c, \theta) = \begin{bmatrix}\sqrt{c} \cos(\theta) \\ \sqrt{c} \sin(\theta)\end{bmatrix} : \theta \in [0, 2\pi)\} $$

The gradient is:
$$ \nabla_\bfx f(\bfx)^\top = 2\bfx^\top 
= 2 \begin{bmatrix} \sqrt{c} \cos(\theta) & \sqrt{c} \sin(\theta)\end{bmatrix}$$

and 

The derivative of curve with respect to $\theta$ the tangent to the curve:
$$\calJ_\theta \bfg(c, \theta) 
= \begin{bmatrix} \frac{\p }{\p \theta} \sqrt{c} \cos(\theta) \\ \frac{\p }{\p \theta} \sqrt{c} \sin(\theta)\end{bmatrix} 
= \begin{bmatrix}-\sqrt{c} \sin(\theta) \\ \sqrt{c} \cos(\theta)\end{bmatrix}$$

$$ \nabla_\bfx f(\bfx)^\top\calJ_\theta \bfg(c, \theta) 
= 2 \begin{bmatrix} \sqrt{c} \cos(\theta) & \sqrt{c} \sin(\theta)\end{bmatrix}
\begin{bmatrix}-\sqrt{c} \sin(\theta) \\ \sqrt{c} \cos(\theta)\end{bmatrix} = 0$$

### Taylor series approximation 
#### A bit more about the step size/learning rate

Recall the Taylor series expansion of a function $f(x)$ around $x_0$

$$f(x) = f(x_0) + \frac{d f(x)}{d x} (x - x_0) + \frac{1}{2!} \frac{d^2 f(x)}{d x^2} (x - x_0)^2 + \dots + \frac{1}{n!} \frac{d^n f(x)}{d x^n} (x - x_0)^n + \dots $$

Vectorized Taylor series expansion is

$$f(\bfx) = f(\bfx_0) + \nabla_\bfx^\top f(\bfx) (\bfx - \bfx_0) + \frac{1}{2!}  (\bfx - \bfx_0)^\top \calH f(\bfx) (\bfx - \bfx_0)^\top + \dots \infty$$

While optimizing around the point $\bfx_t$, we can use the Taylor series expansion to find a local quadratic approximation to find the next best minima:

$$\hat{f}(\bfx_{t+1}) = f(\bfx_t) + \nabla_\bfx^\top f(\bfx) (\bfx_{t+1} - \bfx_t) + \frac{1}{2!}  (\bfx_{t+1} - \bfx_t)^\top \calH f(\bfx) (\bfx_{t+1} - \bfx_t)^\top $$

At the optimal point of the quadratic approximation the derivative is zero:
$$\frac{\p \hat{f}(\bfx_{t+1})}{\p \bfx_{t+1}} = \mathbf{0}^\top$$
$$\nabla_\bfx^\top f(\bfx) + (\bfx_{t+1} - \bfx_t)^\top \calH  f(\bfx) = \mathbf{0}^\top$$

Taking transpose and rearranging the terms we get:
$$ \bfx_{t+1} - \bfx_t = -[\calH  f(\bfx)]^{-1}\nabla_\bfx f(\bfx) $$
$$ \bfx_{t+1} = \bfx_t -[\calH  f(\bfx)]^{-1}\nabla_\bfx f(\bfx) $$

This is the update rule for the Newton's method for optimization. Note that the step size here is inversely proportional the second derivative.

### Taylor series approximation

Approximate the following function to a quadratic function near the point $\bfx_0 = [-2, 3]$

$f (x) = 0.06\exp( 2x_1 +x_2) + 0.05\exp(x_1−2 x_2) + \exp(−x_1) $


$$ f(\bfx) = 0.06 \exp([2, 1]\bfx) + 0.05 \exp([1, -2]\bfx) + \exp([-1, 0]\bfx)$$

$\newcommand{\bfb}{\mathbf{b}}$
Let $\bfb_1^\top = [2, 1]$, $\bfb_2^\top = [1, -2]$, $\bfb_3^\top = [-1, 0]$.

$$\nabla_\bfx^\top f(\bfx) = 0.06 \bfb_1^\top \exp(\bfb_1^\top\bfx) + 0.05 \bfb_2^\top\exp(\bfb_2^\top\bfx) + \bfb_3^\top\exp(\bfb_3^\top\bfx)$$

$$\nabla_\bfx f(\bfx) = 0.06 \bfb_1 \exp(\bfb_1^\top\bfx) + 0.05 \bfb_2 \exp(\bfb_2^\top\bfx) + \bfb_3\exp(\bfb_3^\top\bfx)$$

$$H f(\bfx) = \nabla_\bfx^\top (\nabla_\bfx f(\bfx)) = 0.06 \bfb_1 \exp(\bfb_1^\top\bfx) \bfb_1^\top + 0.05 \bfb_2 \exp(\bfb_2^\top\bfx) \bfb_2^\top + \bfb_3\exp(\bfb_3^\top\bfx)\bfb_3^\top $$

$$H f(\bfx) =  0.06 \exp(\bfb_1^\top\bfx) \bfb_1 \bfb_1^\top + 0.05 \exp(\bfb_2^\top\bfx) \bfb_2 \bfb_2^\top + \exp(\bfb_3^\top\bfx)\bfb_3\bfb_3^\top $$

```python
# Define the function
def f(x):
    
    """
    1. For an input x of shape x.shape = (2,) 
        f(x) must return a scalar
        
    2. For an input x of shape x.shape = (m, 2), 
        f(x) must return an array of shape (m,) 
        which contains f(x) is computed for m values of x

    3. For an input x of shape x.shape = (m, n, 2)
        f(x) must return an array of shape (m, n) 
        which contains f(x) is computed for (m x n) values of x
    """
    return (0.06 * np.exp(x @ [2, 1]) 
            + 0.05* np.exp(x @ [1, -2]) 
            + np.exp(x @ [-1, 0]))

# Compute its derivative, the gradient function
def grad_f(x):     
    """
    1. For an input x of shape x.shape = (2,) 
        grad_f(x) must return a n array of shape (2,)
        
    2. For an input x of shape x.shape = (m, 2), 
        grad_f(x) must return an array of shape (m, 2) 
        which contains grad_f(x) is computed for m values of x

    3. For an input x of shape x.shape = (m, n, 2)
        grad_f(x) must return an array of shape (m, n, 2) 
        which contains grad_f(x) is computed for (m x n) values of x
    """
    coeff1 = np.array([2, 1])
    coeff2 = np.array([1, -2])
    coeff3 = np.array([-1, 0])
    # Slicing using np.newaxis or None, increases the dimension by 1.
    # https://numpy.org/doc/stable/reference/constants.html#numpy.newaxis
    return (0.06 * np.exp(x @ coeff1)[..., None] * coeff1
            + 0.05 * np.exp(x @ coeff2)[..., None] * coeff2
            + np.exp(x @ coeff3)[..., None] * coeff3)

def numerical_jacobian(f, x, h=1e-10):
    n = x.shape[-1]
    eye = np.eye(n)
    x_plus_dx = x + h * eye # n x n
    num_jac = (f(x_plus_dx) - f(x)) / h # limit definition of the formula # n x m
    if num_jac.ndim >= 2:
        num_jac = num_jac.swapaxes(-1, -2) # m x n
    return num_jac
    
# Compare our grad_f with numerical gradient
def check_numerical_jacobian(f, jac_f,  nD=2, **kwargs):
    x = np.random.rand(nD)
    num_jac = numerical_jacobian(f, x, **kwargs)
    return np.allclose(num_jac, jac_f(x), atol=1e-06, rtol=1e-4) # m x n

## Throw error if grad_f is wrong
assert check_numerical_jacobian(f, grad_f)

## Gradient of gradient
def hessian_f(x):
    """
    1. For an input x of shape x.shape = (2,) 
        hessian_f(x) must return a n array of shape (2, 2)
        
    2. For an input x of shape x.shape = (m, 2), 
        hessian_f(x) must return an array of shape (m, 2, 2) 
        which contains hessian_f(x) is computed for m values of x

    3. For an input x of shape x.shape = (m, n, 2)
        hessian_f(x) must return an array of shape (m, n, 2, 2) 
        which contains hessian_f(x) is computed for (m x n) values of x 
    """
    coeff1 = np.array([2, 1])
    coeff2 = np.array([1, -2])
    coeff3 = np.array([-1, 0])
    return (0.06 * np.exp(x @ coeff1)[..., None, None] * np.outer(coeff1, coeff1)
            + 0.05 * np.exp(x @ coeff2)[..., None, None] * np.outer(coeff2, coeff2)
            +  np.exp(x @ coeff3)[..., None, None] * np.outer(coeff3, coeff3))

## Throw error if hessian_f is wrong
assert check_numerical_jacobian(grad_f, hessian_f)


def taylor_series_quad_approx(x0, func, grad_func, hessian_func):
    def quad_func(x):
        x_min_x0 = (x-x0)[..., None] # make column vectors
        x_min_x0_T = (x-x0)[..., None, :] # make row vectors
        grad_f_x0_T = grad_func(x0)[..., None, :] # make row vectors
        return (func(x0) 
                + grad_f_x0_T @ x_min_x0
                + x_min_x0_T @ hessian_func(x0) @ x_min_x0).squeeze(axis=(-1,-2))
            
    return quad_func

def plot_contours(func, ax=None, cmap='Blues_r', levels=20, 
                  xrange=slice(-3,3,21j),
                  yrange=slice(-3,3,21j)):
    x, y = np.mgrid[xrange,
                    yrange]
    bfx = np.concatenate([x[..., None],
                          y[..., None]], axis=-1)
    f = func(bfx)
    if ax is None:
        fig, ax = plt.subplots()
    ctr = ax.contour(x, y, np.log(f), levels, cmap=cmap)
    ax.clabel(ctr, ctr.levels, inline=True, fontsize=6)
    ax.set_xlabel('$x$')
    ax.set_ylabel('$y$')
    return ax
    
# Fit a quadratic curve around this curve:
ax = plot_contours(f, 
                   levels=[0.01, 0.5, 0.8, 1.0, 1.2, 1.6, 1.8, 2.0])
x0 = np.array([0, 2])
ax.plot([x0[0]], [x0[1]], 'ro')
ax.text(x0[0], x0[1], '$x_0$')
quad_func = taylor_series_quad_approx(x0, f, grad_f, hessian_f)
# x, y = np.mgrid[-3:3:21j,
#                     -3:3:21j]
# bfx = np.concatenate([x[..., None],
#                       y[..., None]], axis=-1)
# print(bfx.shape)
# print(f(bfx).shape)
# print(grad_f(bfx).shape)
# print(quad_func(bfx).shape)
plot_contours(quad_func, ax=ax, cmap='Reds_r', 
              levels=[0.1, 0.2, 0.5, 0.8, 1.0, 1.5],
              xrange=slice(-1,2,21j),
              yrange=slice(-2,3,21j))
ax.set_title("Taylor series quadratic fit in red to the blue curve")
plt.show()
```

#### Example Taylor series 1

Following the example above, fit a quadratic function to the function using Taylor series expansion near the points $\bfx_0 = \begin{bmatrix} -1 \\ 1\end{bmatrix}$ then visualize the contour plots of the original function and the quadratic function (50 marks),

$$f(\bfx) = x_1 \exp(-(x_1^2 + x_2^2)) $$

<!-- grade_id: cell-4dccb852958c96b5, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

## Minimization by gradient descent

#### Example 1 : minimization by gradient descent

Find the minimizer of $f (\bfx) = 0.06\exp( 2x_1 +x_2) + 0.05\exp(x_1−2 x_2) + \exp(−x_1) $.

In vector form we can write it as:

$$f (\bfx) = 0.06\exp( [2, 1] \bfx ) + 0.05\exp([1, -2] \bfx) + \exp([-1, 0] \bfx) $$
$$ \nabla_\bfx^\top f(\bfx) = 0.06\exp( [2, 1] \bfx )[2, 1] + 0.05\exp([1, -2] \bfx)[1, -2] + \exp([-1, 0] \bfx) [-1, 0] $$

Algorithm 9.3 Gradient descent method.

**given** a starting point $\bfx \in \text{dom}{f}$ .

**repeat**
1. Choose step size $\alpha_t$ 
2. Update. $\bfx := \bfx  - \alpha_t \nabla f(\bfx) $

**until** stopping criterion is satisfied.

```python
# Define the function
def f(x):
    return (0.06 * np.exp(x @ [2, 1]) 
            + 0.05* np.exp(x @ [1, -2]) 
            + np.exp(x @ [-1, 0]))

# Compute its derivative, the gradient function
def grad_f(x):    
    coeff1 = np.array([2, 1])
    coeff2 = np.array([1, -2])
    coeff3 = np.array([-1, 0])
    # Slicing using None, increases the dimension by 1.
    # 
    return (0.06 * np.exp(x @ coeff1)[..., None] * coeff1
            + 0.05 * np.exp(x @ coeff2)[..., None] * coeff2
            + np.exp(x @ coeff3)[..., None] * coeff3)


assert check_numerical_jacobian(f, grad_f)
f(np.zeros((2,))), grad_f(np.zeros((2,))) 
```

```python
def plot_gradients(func, gradfunc):
    x, y = np.mgrid[-3:3:21j,
                    -3:3:21j]
    bfx = np.concatenate((x[..., None], y[..., None]), axis=-1)
    f = func(bfx)
    dfdx = gradfunc(bfx)
    fig, ax = plt.subplots()
    ctr = ax.contour(x, y, np.log(f), 20, cmap='Blues_r')
    ax.quiver(bfx[..., 0], bfx[..., 1], dfdx[..., 0], dfdx[..., 1])
    ax.set_xlabel('$x$')
    ax.set_ylabel('$y$')
    ax.axis('equal')
    plt.show()
    
```

```python
plot_gradients(f, grad_f)
```

```python
# Implement the gradient descent algorithm
def minimize(x0, f, grad_func, alpha_t=0.2, maxiter=100):
    t = 0
    xt = x0
    grad_f_t = grad_func(xt)
    list_of_xts, list_of_fs = [xt], [f(xt)] # for logging
    while np.linalg.norm(grad_f_t) > 1e-4: # <-- Check for convergence
        xt = xt - alpha_t * grad_f_t # <-- Main update step
        grad_f_t = grad_func(xt) # Compute the next gradient
        
        if t >= maxiter:  # Failsafe, if the algorithm does not converge
            break
        else:
            t += 1
        list_of_xts.append(xt) # for logging
        list_of_fs.append(f(xt)) # for logging
    
    return xt, list_of_xts, list_of_fs
```

```python
x0 = np.array([-2, 2])
#x0 = np.random.rand(2,2) * 4 - 2
OPTIMAL_X, list_of_xts, list_of_fs = minimize(x0,
                                              f, grad_f, 
                                              alpha_t=0.2,
                                              maxiter=1000)

fig, ax = plt.subplots()
ax.plot(list_of_fs)
ax.set_xlabel('t')
ax.set_ylabel('loss')
plt.show()
```

```python
from matplotlib import animation, rc
rc('animation', html='jshtml')

class Anim:
    def __init__(self, fig, ax, func):
        self.fig = fig
        self.ax = ax
        x, y = np.mgrid[-3:3:21j,
                        -3:3:21j]
        bfx = np.concatenate((x[..., None], y[..., None]), axis=-1)
        f = func(bfx)
        self.ctr = self.ax.contour(x, y, np.log(f), 20, cmap='Blues_r')
        self.ax.set_xlabel('x')
        self.ax.set_ylabel('y')
        #self.ax.clabel(self.ctr, self.ctr.levels, inline=True, fontsize=6)
        self.list_of_xs = []
        self.list_of_ys = []
        self.line2, = self.ax.plot([], [], 'r*-')

        
    def anim_init(self):
        return (self.line2,)
        
    def update(self, xt):
        self.list_of_xs.append(xt[0])
        self.list_of_ys.append(xt[1])
        self.line2.set_data(self.list_of_xs, self.list_of_ys)
        return self.line2,
    
fig, ax = plt.subplots()        
a = Anim(fig, ax, f)
animation.FuncAnimation(fig, a.update, frames=list_of_xts[::5],
                        init_func=a.anim_init, blit=True, repeat=False)
```

```python
# The learning rate can be sensitive to the starting points. 
# Observe the behavior for different starting points
# For different starting points you might need different learning rate scheme

x0 = np.random.rand(2) * 6 - 3
OPTIMAL_X, list_of_xts, list_of_fs = minimize(x0,
                                              f, grad_f, 
                                              alpha_t=0.2,
                                              maxiter=200)
fig, ax = plt.subplots()        
a = Anim(fig, ax, f)
animation.FuncAnimation(fig, a.update, frames=list_of_xts[::10],
                        init_func=a.anim_init, blit=True, repeat=False)
```

---

#### Homework 4: Problem 3 (20 marks)

Following the example above, *implement your own* gradient descent algorithm that minimizes the following function (50 marks),

$$f(\bfx) = x_1 \exp(-(x_1^2 + x_2^2)) $$

Test your algorithm with the starting points of $\bfx_0 = \begin{bmatrix} -1 \\ 1\end{bmatrix}$ and $\bfx_0 = \begin{bmatrix} -1 \\ -1\end{bmatrix}$ and learning rate of $\alpha_t = 0.25$. 

---

<!-- grade_id: cell-62edd01f5331c45a, points: 0, grade: false, solution: true -->
```python
x0 = np.array([-1, 1])


def f(x):
    # YOUR CODE HERE
    raise NotImplementedError()

def grad_f(x):
    # YOUR CODE HERE
    raise NotImplementedError()


## Throw error if grad_f is wrong
assert check_numerical_jacobian(f, grad_f)

# Implement the gradient descent algorithm
def minimize(x0, f, grad_func, alpha_t=0.2, maxiter=100):
    t = 0
    xt = x0
    grad_f_t = grad_func(xt)
    list_of_xts, list_of_fs = [xt], [f(xt)] # for logging
    while np.linalg.norm(grad_f_t) > 1e-4: # <-- Check for convergence
        # 1. Update xt using gradient descent update
        # 2. Compute grad_f_t with new xt
        # YOUR CODE HERE
        raise NotImplementedError()
        if t >= maxiter:  # Failsafe, if the algorithm does not converge
            break
        else:
            t += 1
        list_of_xts.append(xt) # for logging
        list_of_fs.append(f(xt)) # for logging
    
    return xt, list_of_xts, list_of_fs


OPTIMAL_X, list_of_xts, list_of_fs = minimize(x0,
                                              f, grad_f, 
                                              alpha_t=0.25,
                                              maxiter=200)

class Anim:
    def __init__(self, fig, ax, func):
        self.fig = fig
        self.ax = ax
        x, y = np.mgrid[-2:1:21j,
                        -1:1:21j]
        bfx = np.concatenate((x[..., None], y[..., None]), axis=-1)
        f = func(bfx)
        self.ctr = self.ax.contour(x, y, f, 
                                   levels=[-4e-1, -3e-1, -2e-1, -1e-1, -7e-2, -5e-2, -1e-2, -1e-3, 0, 1e-3], 
                                   cmap='Blues_r')
        self.ax.set_xlabel('x')
        self.ax.set_ylabel('y')
        #self.ax.clabel(self.ctr, self.ctr.levels, inline=True, fontsize=6)
        self.list_of_xs = []
        self.list_of_ys = []
        self.line2, = self.ax.plot([], [], 'r*-')

        
    def anim_init(self):
        return (self.line2,)
        
    def update(self, xt):
        self.list_of_xs.append(xt[0])
        self.list_of_ys.append(xt[1])
        self.line2.set_data(self.list_of_xs, self.list_of_ys)
        return self.line2,
    
    
fig, ax = plt.subplots()        
a = Anim(fig, ax, f)
anim1 = animation.FuncAnimation(fig, a.update, frames=list_of_xts[::10],
                                init_func=a.anim_init, blit=True, repeat=False)
anim1

```

<!-- grade_id: cell-2802bfb4b2c6edaf, points: 0, grade: true, solution: false -->
```python
assert np.allclose(OPTIMAL_X, [-7.07e-01,  1.06e-04], atol=1e-3, rtol=1e-2)
```

<!-- grade_id: cell-729ffc121c39078d, points: 0, grade: false, solution: true -->
```python
x0 = np.array([-1, -1])
# Repeat minmimization with a different starting point x0
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-e58b107f406f3236, points: 0, grade: true, solution: false -->
```python
assert np.allclose(OPTIMAL_X, [-7.07e-01,  1.06e-04], atol=1e-3, rtol=1e-2)
```

#### Convergence of gradient descent 
(Section 9.3 of [Convex Optimization by Stephen Boyd and Lieven Vandenberghe](https://web.stanford.edu/~boyd/cvxbook/bv_cvxbook.pdf))

```python

```

```python

```