Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All).

Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE", as well as your name and collaborators below:

```python
NAME = ""
COLLABORATORS = ""
```

---

# Homework 3: Problem 6
(10 marks)

$\newcommand{\bfx}{\mathbf{x}}
\newcommand{\bfb}{\mathbf{b}}$
Find the minimum of the following function
$$f(\bfx) = \bfx^\top A^\top A \bfx + 2 \bfb^\top \bfx + c$$

<!-- grade_id: cell-be3d981bec8a49f9, points: 10, grade: true, solution: true -->
YOUR ANSWER HERE

## Homework 3: Problem 7
(10 marks)

$\newcommand{\bfmu}{\boldsymbol{\mu}}$
Find the minimum of the following function (assume $\Sigma_1$ and $\Sigma_2$ to be symmetric and positive semi-definite).
$$f(\bfx) = (\bfx - \bfmu_1)^\top \Sigma_1 (\bfx - \bfmu_1)
+ (\bfx - \bfmu_2)^\top \Sigma_2 (\bfx - \bfmu_2) $$

<!-- grade_id: cell-cb331ee9d8b9f9d8, points: 10, grade: true, solution: true -->
YOUR ANSWER HERE

## Homework 3: Problem 8
(30 marks)

Least square fit in 2D fits a line to a given set of 2D points. This homework asks you to repeat the minimization procedure to find a plane that fits a 2D points. Below you are provided with helper code to load a point cloud, and visualize the points and the plane. Your task is to find the equation of plane that fits the point cloud. 

1. Show the math and derivation for your code
2. Implement the code in python and numpy to find the coefficients of the equation of plane and write your answer in the space provided below. Submit this ipynb file to brightspace. (30 marks)

```python
%pip install plotly requests
def wget(url, filename):
    """
    Download files using requests package
    """
    import requests
    r = requests.get(url)
    with open(filename, 'wb') as fd:
        for chunk in r.iter_content():
            fd.write(chunk)
wget('https://github.com/wecacuee/ECE490-S24-Neural-Networks/raw/master/notebooks/02-linear-models/table_top.npy',
     'table_top.npy')
```

```python
import plotly.graph_objects as go
from plotly.subplots import make_subplots

def points_on_plane(abc):
    a, b, c = abc
    x = np.linspace(-1, 1, 100, dtype=np.float32)
    y = np.linspace(0, 1, 100, dtype=np.float32)

    x, y = np.meshgrid(x, y)
    z = a*x + b*y + c
    return x, y, z

def visualize_plane(pts, abc, fig=None):
    _fig = fig
    fig = make_subplots() if fig is None else fig
    x, y, z = points_on_plane(abc)
    plt_surface = go.Surface(x=x, y=y, z=z)
    
    fig.add_trace(plt_surface)
    visualize_points(pts, fig=fig)
    if _fig is None:
        fig.show()

def visualize_points(pts, fig=None):
    _fig = fig
    fig = make_subplots() if fig is None else fig
    scatter_trace = go.Scatter3d(x=pts[:, 0], y=pts[:, 1], z=pts[:, 2], mode='markers', marker=dict(size=1))
    fig.add_trace(scatter_trace)
    if _fig is None:  
      fig.show()
```

```python
import numpy as np
pts = np.load('table_top.npy').astype('float32') # Nx3
visualize_points(pts)
pts.min(axis=0), pts.max(axis=0)
```

<!-- grade_id: cell-bb7ef5a91325977a, points: 0, grade: false, solution: true -->
```python
def plane_fit(pts):
    # YOUR CODE HERE
    raise NotImplementedError()
    return abc
```

<!-- grade_id: cell-5b7358c1e61a4706, points: 50, grade: true, solution: false -->
```python
# Find abc such that z = a x + b y + c
abc = plane_fit(pts)
visualize_plane(pts, abc)
assert np.abs(pts[:,2] - np.hstack((pts[:, :2], np.ones((pts.shape[0],1)))) @ abc).mean() < 0.4

```

```python

```